> This is a self-correcting activity generated by [nbgrader](https://nbgrader.readthedocs.io). Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Predict diabetes evolution

In this activity, you'll train several regression models to predict the disease progression one year after.

The [Diabetes](https://www4.stat.ncsu.edu/~boos/var.select/diabetes.html) dataset contains ten baseline variables, age, sex, body mass index, average blood pressure, and six blood serum measurements were obtained for each of n = 442 diabetes patients, as well as the response of interest, a quantitative measure of disease progression one year after baseline.

## Environment setup

```python
# Import base packages
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
```

```python
# Import ML packages
import sklearn

print(f"scikit-learn version: {sklearn.__version__}")

from sklearn.datasets import load_diabetes
from sklearn.model_selection import train_test_split, GridSearchCV
from sklearn.metrics import mean_squared_error
```

Output:
```
scikit-learn version: 0.23.1
```

## Step 1: Loading the data

```python
dataset = load_diabetes()

# Put data in a pandas DataFrame
df_diab = pd.DataFrame(dataset.data, columns=dataset.feature_names)
# Add target to DataFrame
df_diab["target"] = dataset.target
# Show 10 random samples
df_diab.sample(n=10)
```

Output:
```
age       sex       bmi        bp        s1        s2        s3  \
403 -0.020045 -0.044642  0.097264 -0.005671 -0.005697 -0.023861 -0.021311   
357 -0.060003  0.050680  0.049840  0.018429 -0.016704 -0.030124 -0.017629   
139  0.023546 -0.044642  0.061696  0.052858 -0.034592 -0.048912 -0.028674   
232  0.012648  0.050680  0.000261 -0.011409  0.039710  0.057245 -0.039719   
227  0.067136  0.050680 -0.029918  0.057449 -0.000193 -0.015719  0.074412   
89  -0.045472 -0.044642 -0.048241 -0.019442 -0.000193 -0.016032  0.067048   
268  0.063504  0.050680  0.088642  0.070073  0.020446  0.037517 -0.050764   
354 -0.023677  0.050680  0.045529  0.021872  0.109883  0.088873  0.000779   
382  0.048974 -0.044642  0.060618 -0.022885 -0.023584 -0.072712 -0.043401   
69   0.016281 -0.044642 -0.046085 -0.005671 -0.075870 -0.061438 -0.013948   

           s4        s5        s6  target  
403 -0.002592  0.061686  0.040343   275.0  
357 -0.002592  0.049769 -0.059067   221.0  
139 -0.002592  0.054724 -0.005
```

## Step 2: Preparing the data

### Question

Split the dataset into training (variables `x_train`, `y_train`) and test sets (variables `x_test`, `y_test`) with a 20% ratio.

<!-- grade_id: cell-4b94ad7a9c319486, points: 0, grade: false, solution: true -->
```python
x = df_diab.drop('target', axis=1)
y = df_diab['target']
x_train, x_test, y_train, y_test = train_test_split(x, y, test_size=0.2)
```

<!-- grade_id: cell-7135eb33d9bae222, points: 1, grade: true, solution: false -->
```python
print(f"x_train: {x_train.shape}. y_train: {y_train.shape}")
print(f"x_test: {x_test.shape}. y_test: {y_test.shape}")

assert x_train.shape == (353, 10)
assert y_train.shape == (353,)
assert x_test.shape == (89, 10)
assert y_test.shape == (89,)
```

Output:
```
x_train: (353, 10). y_train: (353,)
x_test: (89, 10). y_test: (89,)
```

## Step 3: Training several models

```python
def eval_model(model):
    y_train_pred = model.predict(x_train)
    y_test_pred = model.predict(x_test)

    # Train and test MSE
    train_mse = mean_squared_error(y_train, y_train_pred)
    test_mse = mean_squared_error(y_test, y_test_pred)

    print(f"Training MSE: {train_mse:.2f}. Test MSE: {test_mse:.2f}")
    
    return train_mse, test_mse
```

### Question

Create and train a Decision Tree, a MultiLayer Perceptron and a Random Forest on the training data.

Compute their MSE on the training and test data.

<!-- grade_id: cell-1558d1c24e379d78, points: 0, grade: true, solution: true -->
```python
# Import the needed sicki-learn packages
from sklearn.tree import DecisionTreeRegressor
from sklearn.ensemble import RandomForestRegressor
from sklearn.neural_network import MLPRegressor
```

<!-- grade_id: cell-0ee3ffaa00916222, points: 0, grade: false, solution: true -->
```python
# Create and train a Decision Tree

dt_model = DecisionTreeRegressor()
dt_model.fit(x_train, y_train)

eval_model(dt_model)
```

Output:
```
Training MSE: 0.00. Test MSE: 6143.85
(0.0, 6143.853932584269)
```

<!-- grade_id: cell-982fc1c1517d3914, points: 0, grade: false, solution: true -->
```python
# Create and train a MLP

mlp_model = MLPRegressor()
mlp_model.fit(x_train, y_train)

eval_model(mlp_model)
```

Output:
```
Training MSE: 23029.33. Test MSE: 25224.39
/home/jchataigner/anaconda3/lib/python3.8/site-packages/sklearn/neural_network/_multilayer_perceptron.py:582: ConvergenceWarning: Stochastic Optimizer: Maximum iterations (200) reached and the optimization hasn't converged yet.
  warnings.warn(
(23029.334255016955, 25224.387033556686)
```

<!-- grade_id: cell-7d3465c2a61dab34, points: 0, grade: false, solution: true -->
```python
# Create and train a Random Forest

rf_model = RandomForestRegressor()
rf_model.fit(x_train, y_train)
eval_model(rf_model)
```

Output:
```
Training MSE: 487.86. Test MSE: 3296.61
(487.8611722379603, 3296.6079303370784)
```

## Step 4: Tuning the most promising model

### Question

Choose the most promising model and tune it, using a `GridSearchCV` instance stored in the `grid_search_cv` variable.

Your test MSE should be less than 3500.

<!-- grade_id: cell-60f9752b47407958, points: 0, grade: true, solution: true -->
```python
best_model = rf_model
grid_search_cv = GridSearchCV(rf_model, {"max_leaf_nodes": [x for x in range(2,50)], "min_samples_split": [x for x in range(2, 10)]})
```

```python
# Search for the best parameters with the specified classifier on training data
grid_search_cv.fit(x_train, y_train)

# Print the best combination of hyperparameters found
print(grid_search_cv.best_params_)
```

<!-- grade_id: cell-59645e69e9abc0c9, points: 2, grade: true, solution: false -->
```python
# Evaluate best estimator
train_mse, test_mse = eval_model(grid_search_cv.best_estimator_)

assert train_mse < 1000
assert test_mse < 3500
```

```python

```